## Problem: Sliding window

Implement a sliding window for an arbitrary input list. The function should take
the window width and the window increment as inputs and should produce a
sequence of overlapping lists from the input list. For example, given
`x=range(15)`, the following is the output given a window width of `5` and
window increment of `2`.

    [[0, 1, 2, 3, 4],
     [2, 3, 4, 5, 6],
     [4, 5, 6, 7, 8],
     [6, 7, 8, 9, 10],
     [8, 9, 10, 11, 12],
     [10, 11, 12, 13, 14]]

In the event that the input parameters do not yield a complete set of even
sublists, just truncate the ragged tail. For example,

    >>> slide_window(range(18),5,2)
    [[0, 1, 2, 3, 4],
     [2, 3, 4, 5, 6],
     [4, 5, 6, 7, 8],
     [6, 7, 8, 9, 10],
     [8, 9, 10, 11, 12],
     [10, 11, 12, 13, 14],
     [12, 13, 14, 15, 16]]

<!-- grade_id: cell-feb98530bbfed6a6, points: 0, grade: false, solution: true -->
```python
def slide_window(x,width,increment):
# YOUR CODE HERE
    '''
    Implement a sliding window for an arbitrary input list. 
    The function should take the window width and the window increment as inputs 
    and should produce a sequence of overlapping lists from the input list. 
    
    :x :range or list
    :width :int>0
    :increment :int>0
    '''
    from collections import Iterable
    assert isinstance(x, Iterable)
    assert isinstance(width, int)
    assert isinstance(increment, int)
    assert width > 0
    assert increment > 0
    #assert len(x)>=width
    assert width > increment
    if(isinstance(x, set)):
        x = list(x)
    n = 0
    result = []
    groupNumber = int(((len(x) - width)) / increment) + 1
    while(n < groupNumber):
        i = 0
        part = []
        while(i < width):
            index = i + increment * n
            element = x[index]
            part.append(element)
            i = i + 1
        result.append(part)
        n = n + 1
    return result
```

<!-- grade_id: cell-ca9bd3475cce33f2, points: 20, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-21ef5921c0f7a2d6, points: 2, grade: true, solution: false -->
```python
assert slide_window.__doc__ is not None
```

```python

```