Recall the data from last time

    data=[5,4,6,1,9,0,3,9,2,7,10,8,4,7,1,2,7,6,5,2,8,2,0,1,1,1,2,10,6,2]

Write three columns to a comma-separated file:

    data_value, data_value**2, (data_value+data_value**2)/3.

<!-- grade_id: cell-143b614df20a8f3b, points: 0, grade: false, solution: true -->
```python
def write_columns(data,fname):
    '''
    Given data as a list, write three columns to fname.
    
    :param: data 
    :type : list
    :param: fname
    :type: str
    '''
# YOUR CODE HERE

    assert isinstance(data,list) # Verify if the input data is a list
    assert isinstance(fname,str) # Verify if the input fname is a string

    file = open(fname,'w') # Open file in writing mode
    n = 0
    # Use for-loop to write data into the given file
    for element in data:
        assert isinstance(element, int) | isinstance(element, float) # Verify if the input data list elements are integers/floats
        n = n + 1 # the order number of the element we pick-up in the for-loop
        third_col =((element + element**2)/3)
        if(type(element) == int):
            if(n == len(data)):# If element is the last one in the list, just write three result without line break command
                file.write(str(element) + ', ' + str(element**2) + ', ' + '%.3f'%third_col)
            else:
                file.write(str(element) + ', ' + str(element**2) + ', ' + '%.3f'%third_col + '\n')
            
        elif(type(element) == float):
            if(n == len(data)):
                file.write('%.3f'%element + ', ' + '%.3f'%element**2 + ', ' + '%.3f'%third_col)
            else:
                file.write('%.3f'%element + ', ' + '%.3f'%element**2 + ', ' + '%.3f'%third_col + '\n')
            
    file.close() # Close file

```

<!-- grade_id: cell-17cf039687237cdd, points: 10, grade: true, solution: false -->
```python
### BEGIN  TESTS
import numpy as np
from nose.tools import assert_almost_equal
data=[5,4,6,1,9,0,3,9,2,7,10,8,4,7,1,2,7,6,5,2,8,2,0,1,1,1,2,10,6,2]
fname = 'test_data.csv'
write_columns(data,fname)
assert_almost_equal(np.loadtxt(fname,delimiter=',').sum(),1328,2)
    
### END  TESTS
```

```python

```