Download this [corpus of 10,000 common English words](http://localhost:8888/?token=a6840a33f9885242e17dacd3d5ed7114465feb61d54ced90) and do the following:


- What is the longest word?
- What is the longest word that starts with s
- What is the most common starting letter?
- What is the most common ending letter?

<!-- grade_id: cell-e9f6429b2bc8dfbc, points: 0, grade: false, solution: true -->
```python
# you can use this bit of code to download the words from the corpus
from urllib.request import urlopen

u = 'https://storage.googleapis.com/class-notes-181217.appspot.com/google-10000-english-no-swears.txt'
response = urlopen(u)
words = [i.strip().decode('utf8') for i in response.readlines()]

# write a function to compute the longest word
def get_longest_word(words):
# YOUR CODE HERE
    '''
    Given a list, which contains all of the words from 'corpus of 10,000 common English words'.
    To find the longest word.
    
    :param: words
    :type : list
    '''
    assert isinstance(words, list) # Verify if the input data is a list
    assert (len(words)!= 0)        # Verify if the input data has words
    
    # Generate a dictionary ,which contains words and words'length
    dictionary = (dict(zip(words,[len(i) for i in words])))
    #Find the longest word
    longest = max([(i,j) for i,j in dictionary.items()],key = lambda i:i[1])
    res = longest[0]
    return res
    
def get_longest_words_startswith(words, starts):
# YOUR CODE HERE
    '''
    Given a list, which contains all of the words from 'corpus of 10,000 common English words'.
    To find the longest word that start with a particular character.
    
    :param: words
    :type : list
    :param: starts
    :type : str
    '''
    assert isinstance(words, list) # Verify if the input data is a list
    assert (len(words) != 0)       # Verify if the input data has words
    assert(all(words), str)        # Verify if the input word in data list is a string
    assert isinstance(starts, str) # Verify if the input starts is string
    assert starts.isalpha()        # Verify if the input starts is alpha
    assert starts.islower()        # Verify if the input starts is lower string
    assert (len(starts) == 1)      # Verify if the input starts is sigle word
    
    word = []
    res = ''
    # Use for loop to pick up all the words which the start is given word
    for i in words:
        if(i[0] == starts):
            word.append(i)
    
    sorted_words = sorted(word)# Sort the generated list
    # Use for loop to find the longest word starting in given word 
    for i in sorted_words:
        if (len(i) >= len(res)):
            res = i
    return res

def get_most_common_start(words):
# YOUR CODE HERE
    '''
    Given a list, which contains all of the words from 'corpus of 10,000 common English words'.
    To find the most common starting letter.
    
    :param: words
    :type : list
    '''
    assert isinstance(words, list) # Verify if the input data is a list
    assert(all(words), str)        # Verify if the input word in data list is a string 
    assert (len(words) != 0)       # Verify if the input data has words
    
    # Pick up every words' starting letter
    initial = [i[0] for i in words]
    initial_count = {}
    # Use for loop to build up a dictionary containing the beginning letter and the number of occurrences
    for i in initial:
        initial_count[i] = initial.count(i)
    # Find the most common starting letter
    common = max([(i,j) for i,j in initial_count.items()],key = lambda i:i[1])
    res = common[0]
    return res

def get_most_common_end(words):
# YOUR CODE HERE
    '''
    Given a list, which contains all of the words from 'corpus of 10,000 common English words'.
    To find the most common ending letter.
    
    :param: words
    :type : list
    '''
    assert isinstance(words, list) # Verify if the input data is a list
    assert(all(words), str)        # Verify if the input word in data list is a string
    assert (len(words) != 0)       # Verify if the input data has words
    
    # Pick up every words' ending letter
    end = [i[-1] for i in words]
    end_count={}
    # Use for loop to build up a dictionary containing the ending letter and the number of occurrences
    for i in end:
        end_count[i] = end.count(i)
    # Find the most common ending letter
    common = max([(i,j) for i,j in end_count.items()],key = lambda i:i[1])
    res = common[0]
    return res
```

Output:
```
<ipython-input-1-1c2690334855>:41: SyntaxWarning: assertion is always true, perhaps remove parentheses?
  assert(all(words), str)        # Verify if the input word in data list is a string
<ipython-input-1-1c2690334855>:71: SyntaxWarning: assertion is always true, perhaps remove parentheses?
  assert(all(words), str)        # Verify if the input word in data list is a string
<ipython-input-1-1c2690334855>:95: SyntaxWarning: assertion is always true, perhaps remove parentheses?
  assert(all(words), str)        # Verify if the input word in data list is a string
```

<!-- grade_id: cell-9b9db364a733cab0, points: 10, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-7559f8c989e26885, points: 10, grade: true, solution: false -->
```python
### BEGIN  TESTS
assert get_longest_words_startswith(words,'s') in  ['sustainability','simultaneously', 
'specifications', 
'sustainability', 
'superintendent' ]
### END  TESTS
```

<!-- grade_id: cell-3cf6876fc57af686, points: 10, grade: true, solution: false -->
```python
### BEGIN  TESTS
assert get_most_common_start(words)=='c'
### END  TESTS
```

<!-- grade_id: cell-f1359f41f9db52bd, points: 10, grade: true, solution: false -->
```python
### BEGIN  TESTS
assert get_most_common_end(words)=='s'
### END  TESTS
```

```python

```